<problem>
  <script type="loncapa/python">

import json
def vglcfn(e, ans):
    '''
    par is a dictionary that contains two keys, "answer" and "state".
    The value of "answer" is the JSON string that "getGrade" returns.
    The value of "state" is the JSON string that "getState" returns.
    Clicking either "Check" or "Save" registers the current state.

    '''
    par = json.loads(ans)
    # You can use the value of the answer key to grade:
    answer = json.loads(par["answer"])
    print(answer)
    return answer

  </script>
  <customresponse cfn="vglcfn">
    <jsinput gradefn="Survey.getGrade"
	     get_statefn="Survey.getState"
	     set_statefn="Survey.setState"
	     initial_state="{&quot;complete&quot;: false}"
	     html_file="/static/survey.html"
	     width="100%"
	     height="600"
	     sop="true"/>
  </customresponse>
</problem>
